---
title: Test
description: >-
  Test




--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:5cf78fa809
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:ad2c46ca96
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```
